<a href="https://colab.research.google.com/github/Usgupta/FireFightingBot/blob/main/Week08_Cohort.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 8 Problem Set

```python
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import numpy as np
```

## Cohort Sessions

**CS1.** *Reading Data:* Read CSV file for Boston Housing prices.

- **Task 1:** Read the data set. Hint:
    - [Pandas read_csv](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html)
    - Boston's housing data set (filename: `housing_processed.csv`): 
    - [Boston's housing data description](https://www.kaggle.com/vikrishnan/boston-house-prices).

```python
# Task 1
# read CSV file, replace the None
df = None

###
### YOUR CODE HERE
###

display(df)

```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert df.shape == (506, 14)
assert df.columns[0] == 'CRIM' and df.columns[-1] == 'MEDV'
```

- **Task 2:** Display the number of rows and columns. Hint:
    - you can use `df.shape` to get the number of rows and columns


```python
# Task 2
# get the shape from the data frame, replace the None
shape = None

# use the 'shape' variable to get the row and the column
# replace the None
row = None
col = None

###
### YOUR CODE HERE
###

print(shape)
print(row, col)


```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
assert shape == (506, 14)
assert row == 506
assert col == 14
```

- **Task 3:** Display the name of all the columns. Hint:
    - you can use `df.columns` to get the name of all the columns
    - check the meaning of each column using the link above


```python
# Task 3
# display column names, replace the None
names = None

###
### YOUR CODE HERE
###

display(names)

```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert isinstance(names, pd.Index)
assert np.all(names == pd.Index(['CRIM', 'ZN', 'INDUS', 'CHAS', 'NOX', 'RM', 'AGE', 'DIS', 'RAD', 'TAX',
       'PTRATIO', 'B', 'LSTAT', 'MEDV']))
```


- **Task 4:** Do the following:
    - Create a subset data set containing only the following columns: "RM", "DIS", "INDUS" for the features. Make sure it is of `pd.DataFrame` type.
    - Create a subset data set containing only "MEDV" for the target. Make sure it is of `pd.DataFrame` type.


```python
# Task 4
# Specify the columns you want to extract into a list
# replace the None
columns = [None]

# extract the respective columns from the data frame
# replace the None
df_feature = None
df_target = None

###
### YOUR CODE HERE
###

display(df_feature)
display(df_target)


```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_feature, pd.DataFrame)
assert isinstance(df_target, pd.DataFrame)
assert df_feature.shape == (506, 3)
assert df_target.shape == (506, 1)
assert np.all(df_feature.columns == pd.Index(['RM', 'DIS', 'INDUS']))
assert df_target.columns == pd.Index(['MEDV'])
```

**CS2.** *Data Frame Operation:* 

Reference:
- [Indexing and Selecting Data](https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html)

Create separate and new data frame for the columns: `"RM", "DIS", "INDUS", "MEDV"` that satisfies each of the following condition:

- **Task 1:** All records with weighted distances to ﬁve Boston employment centers between 0 to 3.


```python
# specify the columns of interest, replace the None
columns = [None]

# use conditions for row selector and column selector
# replace the None
df_1 = None

###
### YOUR CODE HERE
###

myplot = sns.pairplot(data=df_1)
```

- **Task 2:** All records with average number of room between 5 to 8.


```python
# specify the columns of interest, replace the None
columns = [None]

# use conditions for row selector and column selector
# replace the None
df_2 = None

###
### YOUR CODE HERE
###

myplot = sns.pairplot(data=df_2)
```

- **Task 3:** The first 15 records in the table.


```python
# specify the columns of interest, replace the None
columns = [None]

# use conditions for row selector and column selector
# replace the None
df_3 = None

###
### YOUR CODE HERE
###

display(df_3)

```

- **Task 4:** The last 15 records in the table.


```python
# specify the columns of interest, replace the None
columns = [None]

# use conditions for row selector and column selector
# replace the None
df_4 = None

###
### YOUR CODE HERE
###

display(df_4)
```

- **Task 5:** All records with even index numbers, i.e. index 0, 2, 4, ... .


```python
# specify the columns of interest, replace the None
columns = [None]

# use conditions for row selector and column selector
# replace the None
df_5 = None

###
### YOUR CODE HERE
###

display(df_5)
```

**CS3.** *Histogram and Box plot:* Plot the histogram for the median value in \$1000 for the Boston's housing price. 


Reference:
- [Histogram](https://seaborn.pydata.org/generated/seaborn.histplot.html)
- [Box plot](https://seaborn.pydata.org/generated/seaborn.boxplot.html)

- **Task 1:** Plot the histogram with default bin values.

```python
# plot histogram for MEDV, replace the None
myplot = None

# set the x label, write the code below

# set the y label, write the code below

###
### YOUR CODE HERE
###

```

- **Task 2:** Plot the histogram with 5 bins only. Hint:


```python
# plot histogram for MEDV, replace the None
myplot = None

# set the x label, write the code below

# set the y label, write the code below

###
### YOUR CODE HERE
###

```

- **Task 3:** Plot the histogram with the following bin edges `0, 10, 20, 30, 40, 50`.

```python
# plot histogram for MEDV, replace the None
myplot = None

# set the x label, write the code below

# set the y label, write the code below

###
### YOUR CODE HERE
###

```

- **Task 4:** Plot the same data using a box plot in a horizontal manner.

```python
# plot boxplot for MEDV, replace the None
myplot = None

# set the x label, write the code below

###
### YOUR CODE HERE
###

```

**CS4.** *Scatter plot:* Do the following plots. 

- **Task 1:** Display scatter plot of "RM" versus "MEDV". Hint: 
    - Use Seaborn default theme instead of matplotlib
    - [Seaborn scatter plot](https://seaborn.pydata.org/generated/seaborn.scatterplot.html)
    - Use "RM" as x data
    - Use "MEDV" as y data


```python
# Set the theme parameter to Seaborn's default
sns.set()
```

```python
# Task 5
# plot scatter plot, replace the None 
myplot = sns.scatterplot(x=None, y=None, data=None)

###
### YOUR CODE HERE
###

```

- **Task 2:** Display a scatter plot "weighted distances to ﬁve Boston employment centers" versus "Median value of owner-occupied homes in \$1000s". Use average number of rooms per dwelling as the hue data.


```python
# Task 6
# plot using scatter plot
# replace the None
myplot = None

###
### YOUR CODE HERE
###

```

- **Task 3:** Display a scatter plot "proportion of non-retail business acres per town" versus "Median value of owner-occupied homes in \$1000s". Use "proportion of residential land zoned for lots over 25,000 sq.ft." as the hue data.


```python
# Task 3
# plot using scatter plot and use "hue"
# replace the None
myplot = sns.scatterplot(x=None, y=None, hue="ZN", data=None)

###
### YOUR CODE HERE
###

```

Answer the following:
- How many columns are there?
- How many rows are there?
- What are the names of the columns
- What is the range of the median value in \$1000 when the average number of room per dwelling is 8?
- What's the relationship between distance to five Boston's employment centres and the median value in \$1000?
- What's the relationship between the proportion of non-retail business acres per town adn the median value in \$1000?

**CS5.** *Standardization:* Write a function that takes in data frame where all the column are the features and normalize each column according to the following formula.

$$normalized = \frac{data - \mu}{\sigma}$$

where $\mu$ is the mean of the data and $\sigma$ is the standard deviation of the data. **You need to normalize for each column respectively**. The function should return a new data frame.

Use the following functions from Pandas:
- `df.mean(axis=0)`: This is to calculate the mean along the index axis.
- `df.std(axis=0)`: This is to calculate the standard deviation along the index axis.

Note: Your function should be able to handle a numpy array as well as Panda's data frame. *Hint: use axis=0 argument when calculating mean and standard deviation.*

```python
def normalize_z(dfin):
    ###
    ### YOUR CODE HERE
    ###
    return dfout

```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
data_norm = normalize_z(df_feature)
stats = data_norm.describe()
display(stats)
assert np.isclose(stats.loc["mean", "RM"], 0.0) and \
       np.isclose(stats.loc["std", "RM"], 1.0, atol=1e-3)
assert np.isclose(stats.loc["mean", "DIS"], 0.0) and \
       np.isclose(stats.loc["std", "DIS"], 1.0, atol=1e-3)
assert np.isclose(stats.loc["mean", "INDUS"], 0.0) and \
       np.isclose(stats.loc["std", "INDUS"], 1.0, atol=1e-3)
```

```python
data_norm = normalize_z(df_feature.to_numpy())
assert np.isclose(data_norm[:,0].mean(), 0.0) and \
       np.isclose(data_norm[:,0].std(), 1.0, atol=1e-3)
assert np.isclose(data_norm[:,1].mean(), 0.0) and \
       np.isclose(data_norm[:,1].std(), 1.0, atol=1e-3)
assert np.isclose(data_norm[:,2].mean(), 0.0) and \
       np.isclose(data_norm[:,2].std(), 1.0, atol=1e-3)
```

**CS6.** *Splitting Data Randomly:* Create a function to split the Data Frame randomly. The function should have the following arguments:
- `df_feature`: which is the data frame for the features.
- `df_target`: which is the data frame for the target.
- `random_state`: which is the seed used to split randomly.
- `test_size`: which is the fraction for the test data set (0 to 1), by default is set to 0.5

The output of the function is a tuple of four items:
- `df_feature_train`: which is the train set for the features data frame
- `df_feature_test`: which is the test set for the features data frame
- `df_target_train`: which is the train set for the target data frame
- `df_target_test`: which is the test set for the target data frame

Hint:
- Use [`numpy.random.choice()`](https://docs.scipy.org/doc/numpy-1.13.0/reference/generated/numpy.random.choice.html)


```python
def split_data(df_feature, df_target, random_state=None, test_size=0.5):
    ###
    ### YOUR CODE HERE
    ###
    return df_feature_train, df_feature_test, df_target_train, df_target_test
    

```

```python
df_feature_train, df_feature_test, df_target_train, df_target_test = split_data(df_feature, df_target, random_state=100, test_size=0.3)

display(df_feature_train.describe())
display(df_feature_test.describe())
display(df_target_train.describe())
display(df_target_test.describe())

assert np.all(df_feature_train.count() == 355)
assert np.all(df_feature_test.count() == 151)
assert np.all(df_target_train.count() == 355)
assert np.all(df_target_test.count() == 151)

assert np.isclose(df_feature_train["RM"].mean(), 6.2968)
assert np.isclose(df_feature_test["DIS"].std(), 2.2369)
assert np.isclose(df_target_train["MEDV"].median(), 21.40)
assert np.isclose(df_target_test["MEDV"].median(), 20.90)

```